#### Bi 410/510 (Fall 2019)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

When you are finished submit this notebook through Canvas.

# <span style="color:steelblue;">Milestone 4a</span>

### <span style="color:steelblue;">GC Content</span>

Implement a function named `gc_content` that computes the GC content of a DNA sequence (the percentage of bases that are either Gs or Cs), rounded to two significant digits. 

**Do not use the `count` method.** Instead, write a single for loop that checks each letter in the input string and updates a counter if a letter is a G or a C. 

Example:
```
>>> gc('GATTACA')
28.57
```

<!-- grade_id: gc_4a, points: 3, grade: true, solution: true -->
```python
def gc(seq):
    # YOUR CODE HERE
    gc_count = 0
    for gc in seq:
        if gc == 'G':
            gc_count = gc_count + 1
        if gc == 'C':
            gc_count = gc_count + 1
    s_length = len(seq)
    
    percentage = (float(gc_count)/s_length)*100
    return round(percentage,2)
```

##### Tests 

Use this code cell to test your function as you write it.

```python
def gc(seq):
    # YOUR CODE HERE
```

We'll use these code cells to check your code.

<!-- grade_id: cell-425fdb370cea0bad, points: 1, grade: true, solution: false -->
```python
assert gc('GATTACA') == 28.57
```

Output:
```
<ipython-input-2-57cb6e090ab0> in <module>
----> 1 assert gc('GATTACA') == 28.57

<ipython-input-1-889c41a7cc53> in gc(seq)
      4     for gc in seq:
      5         if gc == 'G':
----> 6             gc_count = gc_count + 1
      7         if gc == 'C':
      8             gc_count = gc_count + 1

UnboundLocalError: local variable 'gc_count' referenced before assignment
```

<!-- grade_id: cell-4bd98d4a857264c3, points: 1, grade: true, solution: false -->
```python
assert gc('AAACCCGGGTTT') == 50.0
```

<!-- grade_id: cell-7529958e880100b4, points: 1, grade: true, solution: false -->
```python
assert gc('AAATTT') == 0.0
```

### <span style="color:steelblue;">Acronym</span>

Implement a function named `acronym` that will create an acronym from a list of words.  The result should be a single string that has the first character of each word from the list.

Example:
```
>>> acronym(["North", "Atlantic", "Treaty", "Organization"])
'NATO'
```

<!-- grade_id: acronym, points: 3, grade: true, solution: true -->
```python
def acronym(lst):
    # YOUR CODE HERE  
    for x in lst:
        print (''.join(x[0] for x in x.split()))
```

##### Tests 

Use this code cell to test your function as you write it.

```python
def acronym(lst):
    # YOUR CODE HERE  
    returned_string = ''
    for x in lst:
        returned_string += x[0]
    return returned_string
```

We'll use these code cells to check your code.

<!-- grade_id: cell-996686965e92ca7e, points: 1, grade: true, solution: false -->
```python
assert acronym(["North", "Atlantic", "Treaty", "Organization"]) == 'NATO'
```

<!-- grade_id: cell-cdb3e7cce1cffac3, points: 1, grade: true, solution: false -->
```python
assert acronym(["Martin", "Luther", "King"]) == 'MLK'
```

### <span style="color:steelblue;">Birthday Paradox</span>

A fun ice-breaker for a new group of people is to have everyone compare birthdays to see if any two people were born on the same day of the year. It may not seem intuitive, but with only 23 people in a group there is a 50\% chance that two of them will have the same birthday. With a group of 100 the probability that two people will have the same birthday is more than 99.999\%.

To simulate this phenomenon in Python write a function named `paradox`. Create an empty list, and then use a loop that adds numbers to the list. On each iteration:
* choose a random number between 1 and 365 (use the `randint` function from the `random` library)
* if the number is already in the list exit the loop
* otherwise add the number to the list and continue.

The function should return the number of times the loop iterates before finding a duplicate. If it will help, you can print the contents of your list and the random number you draw on each loop iteration. Here is a sample output with this more verbose output:
```
>>> paradox()
look for 343 in []
look for 130 in [343]
look for 271 in [343, 130]
look for 192 in [343, 130, 271]
look for 97 in [343, 130, 271, 192]
look for 322 in [343, 130, 271, 192, 97]
look for 271 in [343, 130, 271, 192, 97, 322]
6
```

<!-- grade_id: paradox, points: 3, grade: true, solution: true -->
```python
from random import randint

def paradox():
    lst = []
    randint(1,365)
   
    x = randint(1,365)
    
    while x not in lst:
        print('look for ' + str(x) + ' in ' + str(lst))
        lst.append(x)
        x = randint(1,365)
        
    #print('look for ' + str(x) + ' in ' + str(lst))
        
    return len(lst)

    #for x in lst:
    
        
```

##### Tests 

Use this code cell to test your function as you write it.

```python
paradox()
```

Output:
```
look for 321 in []
look for 202 in [321]
look for 14 in [321, 202]
look for 62 in [321, 202, 14]
look for 266 in [321, 202, 14, 62]
look for 14 in [321, 202, 14, 62, 266]
5
```

We'll use these code cells to check your code.  Each time you call your function you should get back a number between 1 and 365.

<!-- grade_id: cell-3a430d9c6a27297d, points: 2, grade: true, solution: false -->
```python
for i in range(5):
    assert 1 <= paradox() <= 365
```

### <span style="color:red;">Short Answer Question</span>

Don't forget to complete the short-answer question on the exam page.

Double-click on the markdown cell below and enter your answer in this cell.

<!-- grade_id: cell-cd72593eb7ea9d6f, points: 4, grade: true, solution: true -->
1. loop executed 5 times
2. when ball is set to 30: char is (ball-1)//15 = 29//15 = 1
3. 1 + 30 = 31 (value of message is 31?)