#### Bi 410/510 (Fall 2019)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

When you are finished submit this notebook through Canvas.

# <span style="color:steelblue;">Milestone 5a</span>

### <span style="color:steelblue;">Points Scored in a Basketball Game</span>

Implement a function named `points_scored` that will read a data file (in CSV format) with scoring data from a basketball game.  Each line in the file will have four pieces of data:
* a player's name
* the number of 3-point baskets made by that player
* the number of 2-point baskets made
* the number of free throws (1-point baskets) made

Your function should compute the total number of points scored by the entire team.

As an example, suppose a file named `demo.csv` has these three lines:
```
ron,1,1,1
neville,0,0,5
cho,3,0,3
```
If we pass this file name to the function, it should return 23 (Ron scored 6 points, Neville scored 5, and Cho scored 12):
```
>>> points_scored('demo.csv')
23
```

<!-- grade_id: points_scored, points: 3, grade: true, solution: true -->
```python
def points_scored(fn):
    # YOUR CODE HERE
    total = 0.0
    point = []
    for point in open(fn):
        point = point.split(',')
        
        t1 = point[1]
        t2 = point[2]
        t3 = point[3]     
    
        value = int(t1)*3 + int(t2)*2 + int(t3)*1    
        total +=value
        
    return int(total)
       
```

##### Sandbox 

To test your function you can use `demo.csv`.  The expected result is 23.

```python
points_scored('demo.csv')
```

Output:
```
<ipython-input-2-0969792f44e2> in <module>
----> 1 points_scored('demo.csv')

<ipython-input-1-5f04fa66b5ea> in points_scored(fn)
     10 
     11         value = int(t1)*3 + int(t2)*2 + int(t3)*1
---> 12         total +=value
     13 
     14     return int(total)

UnboundLocalError: local variable 'total' referenced before assignment
```

##### Tests 

<!-- grade_id: cell-2cd8e4a8510203af, points: 2, grade: true, solution: false -->
```python
assert points_scored('game.csv') == 91
```

### <span style="color:steelblue;">Points Dictionary</span>

Define a function named `points_dictionary` that will also be passed the name of a CSV file with basketball game data.  This function should return a dictionary that shows how many points each player scored.  The keys will be player names, and the values will be the number of points scored by that player.

Example:
```
>>> points_dictionary('demo.csv')
{'ron': 6, 'neville': 5, 'cho': 12}
```

<!-- grade_id: points_dict, points: 3, grade: true, solution: true -->
```python
def points_dictionary(fn):
    
    tol_p = 0.0
    array = {}
    
    
    for points in open(fn):
        tol_p = points.strip().split(',')
        
        name = tol_p[0] #player name
        onePointers = tol_p[1] #1 point baskets
        twoPointers = tol_p[2] #2 point baskets
        threePointers = tol_p[3] #3 point baskets
        
        totalPoints = int(onePointers)*3 + int(twoPointers)*2 + int(threePointers)*1
    
        array.update({name : totalPoints})
        
        
    return array
        
```

##### Sandbox 

def points_dictionary(fn):
   
   tol_p = 0.0
    
   for points in open(fn):
        tol_p = points.strip().split(',')

```python
points_dictionary('game.csv')
```

Output:
```
{'sabrina': 28,
 'ruthy': 18,
 'satou': 10,
 'mallory': 9,
 'maite': 5,
 'anneli': 9,
 'aina': 6,
 'oti': 4,
 'lydia': 2,
 'justine': 0,
 'sierra': 0}
```

##### Tests 

<!-- grade_id: cell-c99afc74fa61dfbb, points: 2, grade: true, solution: false -->
```python
g = points_dictionary('game.csv')
assert len(g) == 11 
assert g['sabrina'] == 28
assert sum(g.values()) == 91
```

### <span style="color:steelblue;">Key Generator</span>

The lecture notes for October 30 (`Dictionaries.ipynb` on Canvas) showed how to encrypt a string using a "substitution cipher," a dictionary that maps letters to the corresponding code letter, for example
```
{ ’H’: ’N’, ’E’: ’T’, ...}
```

There are 26 letters in the English alphabet, so this dictionary has 26 entries, one for each letter. 

Define a function named `key_gen` that will create and print a dictionary like this. The function does not have any inputs. Each time you call it the function should make a random cipher and return it.

Example:
```
>>> key_gen()
{’A’: ’Z’, ’B’: ’S’, ’C’: ’B’, ... }
```

Don’t worry about the order of the key-value pairs, Python is allowed to print them in any order. Just make sure you have 26 pairs in all, with one for each letter of the alphabet.

#### Hints and Suggestions 

(1) Use the `ascii_uppercase` string in Python's `string` library.  You can import it with this statement:
```
from string import ascii_uppercase
```

(2) If you want to make a list containing each letter in the alphabet this will do it:
```
alphabet = list(ascii_uppercase)
```
(try it in a "sandbox cell" to make sure it does what you expect).

(3) There is a function named `shuffle` in the `random` library that will scramble a list, reorganizing the items into a random order.  Example:
```
from random import shuffle
>>> shuffle(alphabet)
```
(try that, too).

<!-- grade_id: key_gen, points: 4, grade: true, solution: true -->
```python
from string import ascii_uppercase
from random import shuffle

def key_gen():
    normalAlphabet = list(ascii_uppercase)
    shuffledAlphabet = list(ascii_uppercase)
    
    shuffle(shuffledAlphabet)
    cipher = {}
    
    n = 0    
    for item in normalAlphabet:
        cipher.update({item : shuffledAlphabet[n]})
        n += 1
    return cipher
    
        
    
```

##### Sandbox 

Each time you call your function you should get back a dictionary where the keys are the 26 upper case letters.  The values should also be the 26 letters, but in a random order.  Don't worry if Python prints the dictionary out of order. 

```python
key_gen()
```

Output:
```
{'A': 'X',
 'B': 'W',
 'C': 'Y',
 'D': 'A',
 'E': 'Z',
 'F': 'F',
 'G': 'I',
 'H': 'B',
 'I': 'N',
 'J': 'Q',
 'K': 'M',
 'L': 'G',
 'M': 'S',
 'N': 'P',
 'O': 'K',
 'P': 'R',
 'Q': 'L',
 'R': 'O',
 'S': 'E',
 'T': 'T',
 'U': 'V',
 'V': 'U',
 'W': 'D',
 'X': 'C',
 'Y': 'J',
 'Z': 'H'}
```

##### Tests 

A valid cipher has every letter of the alphabet as a key, and every letter also appears somewhere as a value.  That means the sorted list of keys should be the same as the sorted list of values.

<!-- grade_id: cell-d9e92e4adb474039, points: 2, grade: true, solution: false -->
```python
dct = key_gen()
assert len(dct) == 26
assert sorted(dct.keys()) == sorted(dct.values())
```

### <span style="color:red;">Short Answer Question</span>

Don't forget to complete the short-answer question on the exam page.

Double-click on the markdown cell below and enter your answer in this cell.

<!-- grade_id: cell-0545f0c3cc3f3e4a, points: 4, grade: true, solution: true -->
1. comp is a dictionary and has 4 main pairs (values)
2. Returning res will accomplish the following:
    it will translate A to T, T to A, G to C, C to G in quotation marks (''). AKA: 'GATTACA' -> 'CTAATGT' (This is DNA transcription).