#### Bi 410/510 (Fall 2019)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

As you work on the projects write your Python code in the cells that have a comment that says `"YOUR CODE HERE"` or `"YOUR ANSWER HERE"`.

**IMPORTANT:** Do not copy or delete these cells.

You can add additional "sandbox" cells if you wish, but make sure you don't duplicate a solution cell.

When you are finished submit this notebook through Canvas.

# <span style="color:steelblue;">Milestone 8A</span>

#### <span style="color:steelblue;">Imports</span>

Execute this code cell to import the libraries used for this exam:

```python
import os
import shutil
```

#### <span style="color:steelblue;">Directory</span>

The exercises in this exam are based on a folder named `greeks` that is created by unzipping a file named `greeks.zip`.  Make sure the zip file is in the same directory as this notebook, then execute this code cell to unzip the file: 

```python
! rm -rf greeks
! unzip -qq greeks.zip
```

**Note:**  Some of the code you write for the exam will update the folder.  If you need to restore the folder to its original condition (_e.g._ before running a test) simply execute the code cell again to delete the current folder and unzip a new copy.

### <span style="color:steelblue;">Count Folders</span>

Fill in the definition of the function named `count_folders`.  It will be passed a string that is the path to a directory, and it should return the number of subdirectories found there.

<!-- grade_id: count_folders, points: 3, grade: true, solution: true -->
```python
! rm -rf greeks
! unzip -qq greeks.zip

import os
import shutil

def count_folders(dirname):
    foldersLst = []
    numFol = 0
    
    for name in os.listdir(dirname):
        filepath = os.path.join(dirname, name)
        
        if os.path.isdir(filepath):
            foldersLst.append(name)
    numFol += len(foldersLst)
    print(foldersLst)
    
    print('number of subdirectories within', dirname,':', numFol)
    return(numFol)
```

#### Sandbox 

You can use this code cell to test your function.

```python

```

#### Tests 

<!-- grade_id: cell-c15edd738dd23525, points: 1, grade: true, solution: false -->
```python
assert count_folders('greeks') == 6
```

Output:
```
['alpha', 'beta', 'chi', 'delta', 'epsilon', 'gamma']
number of subdirectories within greeks : 6
```

### <span style="color:steelblue;">Create Subdirectory</span>

Fill in the definition of the function `create_subdirectory` so it adds a new directory to an existing directory.  The parameters passed to the function will be the path to the existing directory and a string that will be the name of the new subdirectory.

The function should check to see if the subdirectory exists already.  If so, it should print an error message and return False.  Otherwise it should create the directory and return True.

Examples:
```
>>> create_subdirectory('greeks', 'consonants')
True

>>> create_subdirectory('greeks','alpha')
directory exists
False
```

<!-- grade_id: create_subdir, points: 3, grade: true, solution: true -->
```python
import os
import shutil

! rm -rf greeks
! unzip -qq greeks.zip

def create_subdirectory(dirname,subname):
    filepath = os.path.join(dirname, subname)
    
    if not os.path.isdir(filepath):
        os.mkdir(filepath)
        print('subdirectory created:', filepath)
        return True
    
    else:
        print('subdirectory already exists:', filepath)
        return False
    
```

#### Sandbox 

```python

```

#### Tests 

These test cells will create two new subdirectories in `greeks`.  The first two calls to `create_subdirectory` should succeed, the third is expected to fail because there is already a folder with that name.

<!-- grade_id: cell-502ea03b0ff7954f, points: 1, grade: true, solution: false -->
```python
assert create_subdirectory('greeks','consonants') and create_subdirectory('greeks','vowels')
```

Output:
```
subdirectory created: greeks/consonants
subdirectory created: greeks/vowels
```

<!-- grade_id: cell-8b2d8e1f349d6592, points: 1, grade: true, solution: false -->
```python
for d in ['consonants', 'vowels']:
    assert d in os.listdir('greeks')
```

<!-- grade_id: cell-218d6a33fc70f9cd, points: 1, grade: true, solution: false -->
```python
assert not create_subdirectory('greeks','alpha')
```

Output:
```
subdirectory already exists: greeks/alpha
```

### <span style="color:steelblue;">Move Files</span>

Fill in the body of the `move_files` function.  It will be passed two directory names.  The function should move all the items in the first directory to the second directory.

Example: to move all the files in `alpha` to `vowels` the call will be
```
>>> move_files(greeks/alpha, greeks/vowels)
```

<!-- grade_id: move_files, points: 3, grade: true, solution: true -->
```python
def move_files(source, destination):
    
    if os.path.isdir(source):
        for file in os.listdir(source):
            shutil.move(source+"/"+file, destination)
            
            print(file, 'from', source, 'moved to:', destination)
```

#### Sandbox 

```python

```

#### Tests 

The first test cell will call the function to move the contents of the `alpha` and `epsilon` folders to `vowels`, then check to see if all the files were moved.

<!-- grade_id: cell-0512dbcf1a8b875d, points: 1, grade: true, solution: false -->
```python
move_files('greeks/alpha', 'greeks/vowels')
move_files('greeks/epsilon', 'greeks/vowels')
assert len(os.listdir('greeks/vowels')) == 5
```

Output:
```
aedesia.txt from greeks/alpha moved to: greeks/vowels
aristotle.txt from greeks/alpha moved to: greeks/vowels
epictetus.txt from greeks/epsilon moved to: greeks/vowels
epicurus.txt from greeks/epsilon moved to: greeks/vowels
euclid.txt from greeks/epsilon moved to: greeks/vowels
```

The second test moves the contents of the other folders to `consonants`.

<!-- grade_id: cell-16cb5e51959e23aa, points: 1, grade: true, solution: false -->
```python
for dirname in ['beta', 'chi', 'delta', 'gamma']:
    source = os.path.join('greeks', dirname)
    move_files(source, 'greeks/consonants')
assert len(os.listdir('greeks/consonants')) == 2
```

Output:
```
basilides.txt from greeks/beta moved to: greeks/consonants
diogenes.txt from greeks/delta moved to: greeks/consonants
```



### <span style="color:steelblue;">Remove Empty Directories</span>

Fill in the body of the function named `delete_empty` so it finds and deletes all the empty subdirectories in a directory.  It should return a list of names of each subdirectory that was deleted.

We'll test the function by calling `delete_empty(greeks)`.  At this point all of the original subdirectories should be empty so they should all be deleted:
```
>>> delete_empty('greeks')
['greeks/beta', 'greeks/chi', 'greeks/delta', 'greeks/gamma', 'greeks/epsilon', 'greeks/alpha']
```

Note: the names in the list can be in any order.

<!-- grade_id: delete_empty, points: 3, grade: true, solution: true -->
```python
def delete_empty(dirname):
    emptyFol = [] #list of which filepaths are empty
    
    for folder in os.listdir(dirname):
        filepath = os.path.join(dirname, folder)
        
        if os.path.isdir(filepath):
            if len(os.listdir(filepath)) == 0:
                emptyFol.append(filepath)
                os.rmdir(filepath)
                print(filepath, 'removed sucessfully!')
    return emptyFol
```

#### Sandbox 

```python

```

#### Test 

The first test should return the list of all the original directories.  The second should return an empty list because there shouldn't be any empty directories left.

<!-- grade_id: cell-230a0e7056359dc6, points: 1, grade: true, solution: false -->
```python
lst = delete_empty('greeks')
assert len(lst) == 6
assert 'greeks/alpha' in lst
```

<!-- grade_id: cell-985cd3c8ca5adff8, points: 1, grade: true, solution: false -->
```python
assert delete_empty('greeks') == []
```